## `lab05`—Scientific Applications with `sympy`

![](./img/header.png)

Objectives

-   Plot expressions using SymPy.
-   Apply scientific modeling across a series of case studies.

### Kepler's Laws

In 1617, Johannes Kepler published _Epitome astronomiae Copernicanae_ (or _Epitome of Copernican Astronomy)_, which introduced his celebrated laws of motion.

[Kepler's three laws of motion](https://en.wikibooks.org/wiki/General_Astronomy/Kepler%27s_Laws) introduce a mathematical model of elliptical (rather than circular) orbits.  They are:

1.  The orbits of the planets are ellipses, with the sun at one focus.
2.  The line between a planet and the sun sweeps out equal areas in the plane of the planet's orbit over equal times.
3.  The time required for a planet to orbit the sun, called its period, is proportional to half the long axis of the ellipse raised to the 3/2 power.

In this portion of the lab, you will carry out calculations in support or demonstration of these laws.

##### Elliptical Orbits

An elliptical orbit has two bodies moving such that the lighter body orbits the heavier body as one focus of an ellipse.

![](https://upload.wikimedia.org/wikipedia/commons/c/c9/Kepler1.gif)

For real planets, the foci are quite close, as planetary orbits are only _slightly_ eccentric.  For comets and other bodies, the foci may lie far apart.

An equation describing the elliptical orbit may be written:

$$
\frac{x^2}{a^2} + \frac{y^2}{b^2}
=
1
$$

with eccentricity $e$ described by $b^2 = a^2\left( 1 - e^2 \right)$.  $a$ and $b$ are the _semimajor_ and _semiminor_ axes, respectively, of the ellipse.

```python
import sympy
x,y,a,b,e = sympy.S( 'x,y,a,b,e' )
```

In preparation for solving and plotting an orbit, we suppose that we are given the eccentricity $e$ and semimajor axis $a$, and need to obtain $b$ from that information.  First, solve for $b$ as `b_soln`:

```python
b_soln = ___

print( b_soln )
```

Next, `subs`titute in the values `e = 0.2056` and `a = 57909050000`, using SI units for convenience.  (These are the values for the planet Mercury ☿.)

```python
e_ = 0.2056
a_ = 57909050000
b_ = b_soln[ 1 ].___

print( 'e = %f; a = %f; b = %f' % ( e_,a_,b_ ) )
```

Now that we know $a$ and $b$, we can solve for the elliptical orbit.

Solving for $y$ in terms of $x$ will lead to multiple expressions since the ellipse is multivalued along the $x$ axis.

```python
y_ = sympy.solve( ___,y )

print( y_ )
```

We plot using a library we haven't seen yet, MatPlotLib.  This is actually the library behind SymPy's plotting ability, and we'll examine it during `num/plotting`.

A true circle is shown in red underlying the eccentric ellipse.

```python
# Plotting boilerplate---leave this as-is.
import matplotlib as mpl
import matplotlib.pyplot as plt
%matplotlib inline
import seaborn as sns
sns.set_context( 'talk',font_scale=1.5,rc={ 'lines.linewidth':5 } )
```

```python
# Plotting code.
plot1 = sympy.plotting.plot_parametric( ( x,y_[ 0 ] ),( x,-1e12,1e12 ),show=False )
plot2 = sympy.plotting.plot_parametric( ( x,y_[ 1 ] ),( x,-1e12,1e12 ),show=False )
plot_circle1 = sympy.plotting.plot_parametric( ( x,+(-x**2 + 3353458071902500000000)**0.5 ),( x,-1e12,1e12 ),show=False,line_color='red' )
plot_circle2 = sympy.plotting.plot_parametric( ( x,-(-x**2 + 3353458071902500000000)**0.5 ),( x,-1e12,1e12 ),show=False,line_color='red' )
plot_circle1.append( plot_circle2[ 0 ] )
plot_circle1.append( plot1[ 0 ] )
plot_circle1.append( plot2[ 0 ] )
mpl.rcParams[ 'figure.figsize' ] = 10,10   # make the figure square, necessary in Jupyter
plot_circle1.show()
```

##### Equal Area in Equal Time

Kepler's second law posits that the orbiting planet's velocity varies such that the same area is swept out during the same time.

![](https://upload.wikimedia.org/wikipedia/commons/7/75/Kepler2.gif)

![](https://upload.wikimedia.org/wikipedia/commons/thumb/6/69/Kepler-second-law.gif/220px-Kepler-second-law.gif)

As you can imagine, this becomes a somewhat complicated integral to write, depending on the term we solve for.  For convenience, we will switch to polar terms in our expression.

$$
\frac{dA}{dt}
=
\frac{1}{2}r^2 \frac{d\theta}{dt}
$$

and thus some calculus and algebra yield

$$
A(t)
=
\int_{\theta_0}^{\theta_1} d\theta\, \frac{r^2}{2}
=
\frac{1}{2} ab \left( \theta_1 - \theta_0 \right)
\text{.}
$$

The largest $t$ is the period of the orbit, and the planet sweeps out the area of the ellipse $\pi ab$ during that time.

$$
A(P)
=
\frac{1}{2} ab \left( 2\pi - 0 \right)
=
\pi ab
$$

Use SymPy to verify that the area traced from $\theta \in \left[ 0,\frac{\pi}{4} \right]$ is the same as that traced from $\theta \in \left[ \frac{3\pi}{8},\frac{5\pi}{8} \right]$.

Since SymPy doesn't like `float`s, prefering rational numbers, use `/2` not `*0.5` when figuring the one-half ratio.

<!-- grade_id: cell-2ada0376f491a468, points: 0, grade: false, solution: true -->
```python
area_1 = ___
area_2 = ___
```

<!-- grade_id: cell-0903030db1d816bc, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from math import isclose
assert ( area_1 == area_2 ) and (( area_1 == sympy.pi*a*b/8 ) or ( area_1 == float(sympy.pi*a*b/8) ))
print( 'Success!' )
```

##### Proportionality of Periods

Given the period of a planet or satellite, one can determine how wide the orbit is, and _vice versa_.

$$
P^2
\propto
a^3
$$

where $P$ is the period in years and $a$ is the length of the semimajor axis in astronomical units (not SI units as above).

Calculate the period of Mars ♂︎ given that its semimajor axis is $a = 1.87 \,\text{AU}$.

<!-- grade_id: cell-ab084336c256f58e, points: 0, grade: false, solution: true -->
```python
P = ___
```

<!-- grade_id: cell-1008d824d54d7721, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
test_P,test_a = sympy.S( 'P,a' )
from math import isclose
assert ( P == test_a.subs( test_a,1.87 ) ** ( 3 / 2 ) ) or ( isclose( P,1.87,rel_tol=1e-2 ) )
print( 'Success!' )
```

Calculate the semimajor axis of Jupiter ♃ given that its period is $11.86 \,\text{yr}$.

<!-- grade_id: cell-72b3e4058171eaaf, points: 0, grade: false, solution: true -->
```python
a = ___
```

<!-- grade_id: cell-839539944c4d5c32, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
test_P,test_a = sympy.S( 'P,a' )
from math import isclose
assert ( a == test_P.subs( test_P,11.86 ) ** ( 2 / 3 ) ) or ( isclose( P,5.2,rel_tol=1e-2 ) )
print( 'Success!' )
```

### Trajectories

A basic ballistic trajectory is subject only to the initial impetus.  More advanced trajectory calculations include thrust during flight and the action of external forces such as gravity.  Drag can be considered among the relevant forces as well.

In this section, we will calculate some trajectories by position and velocity.

##### Simple Trajectory

A simple trajectory describes an object moving only from the initial impulse, such as a rocket in deep space with the engines powered down.  We write the velocity as constant $v_0$, and integrate to obtain the position $x(t)$ as a function of $v_0$.

$$
v(t)
=
v_0
$$

$$
x(t)
=
\int_{\tau=t_0}^{t} d\tau\, v_0
$$

(We will use the Greek equivalent of Roman letters as the variables of integration since $t$ means the current time.)

```python
x,x0,v,v0,tau,t,tf = sympy.S( 'x,x0,v,v0,tau,t,tf' )
v = ___
x = sympy.integrate( v.subs( t,tau ),( tau,0,t ) )   # Note that tau not t is our variable of integration, so we subs it.
```

We plot velocity $v$ in red and position $x$ in blue.

```python
plot_v = sympy.plotting.plot( v.subs( v0,10 ),( t,0,10 ),line_color='red',show=False )
plot_x = sympy.plotting.plot( x.subs( v0,10 ),( t,0,10 ),show=False )
plot_v.append( plot_x[ 0 ] )
plot_v.show()
```

##### Powered Trajectory

That case was basically trivial, so let's try one subject to constant acceleration from a source like an [ion engine]().

$$
v(t)
=
a_0 t
$$

$$
x(t)
=
\int_{\tau=t_0}^{t} d\tau\, a_0 \tau
$$

```python
x,x0,v,a0,tau,t,tf = sympy.S( 'x,x0,v,a0,tau,t,tf' )
v = ___
x = ___  #should be able to write this the same as before
```

Once again, we plot velocity and position for this trajectory:

```python
plot_v = sympy.plotting.plot( v.subs( a0,10 ),( t,0,10 ),line_color='red',show=False )
plot_x = sympy.plotting.plot( x.subs( a0,10 ),( t,0,10 ),show=False )
plot_v.append( plot_x[ 0 ] )
plot_v.show()
```

##### External Forces

Next we consider a scenario with external forces including drag $a_D$ acting on the rocket.  In this case, we posit that acceleration due to drag will vary as a function of time (really, of position, but we're simplifying) as the rocket moves through a thinning atmosphere, $a_D = -\exp\left( 2-t \right)$.  Since gravity is a constant acceleration and we are treating the rocket's thrust as constant, the behavior won't be fundamentally different.

$$
a(t)
=
a_0 + a_D(t)
=
a_0 - \exp\left( 2-t \right)
$$

$$
v(t)
=
\int_{\tau=t_0}^{t} d\tau\, \left( a_0 - \exp\left( 2-\tau \right) \right)
$$

$$
x(t)
=
\int_{\tau=t_0}^{t} d\tau\, v(\tau)
$$

```python
x,x0,v,a0,tau,t,tf = sympy.S( 'x,x0,v,a0,tau,t,tf' )
a = a0 - sympy.exp( 2-t )
v = ___
x = ___  #should be able to write this the same as before
```

Plotting the same:

```python
plot_v = sympy.plotting.plot( v.subs( a0,10 ),( t,0,10 ),line_color='red',show=False )
plot_x = sympy.plotting.plot( x.subs( a0,10 ),( t,0,10 ),show=False )
plot_v.append( plot_x[ 0 ] )
plot_v.show()
```

##### Odd Accelerations

An experimental rocket fuel has been found to burn with a strange sinusoidal acceleration, $a_F = 50 \sin( 10 t )$.  Subject to a standard acceleration $a_0 = -\frac{1}{2}$, will this fuel be able to free the rocket from the planet's gravitational field, or will the ride be too bumpy and ultimately fail?

To answer this question, you'll need to construct $a$ as a function of $a_0$ and $a_F$.

```python
x,x0,v,a0,tau,t,tf = sympy.S( 'x,x0,v,a0,tau,t,tf' )
a = 50*sympy.sin( 10*t ) + a0
v = ___
x = ___  #should be able to write this the same as before
```

Plot the resulting velocity and trajectory:

```python
plot_v = sympy.plotting.plot( v.subs( a0,-0.5 ),( t,0,10 ),line_color='red',show=False )
plot_x = sympy.plotting.plot( x.subs( a0,-0.5 ),( t,0,10 ),show=False )
plot_v.append( plot_x[ 0 ] )
plot_v.show()
```

A different rocket fuel has been found to burn with the acceleration, $a_F = 2 \log( 10 t )$.  Subject to a standard acceleration $a_0 = -\frac{1}{2}$, check this fuel on the same planet to see if the rocket can escape the gravitational field.

To answer this question, you'll need to construct $a$ as a function of $a_0$ and $a_F$.

```python
x,x0,v,a0,tau,t,tf = sympy.S( 'x,x0,v,a0,tau,t,tf' )
a = 2*sympy.log( 10 * t ) + a0
v = ___
x = ___  #should be able to write this the same as before
```

```python
plot_v = sympy.plotting.plot( v.subs( a0,-0.5 ),( t,0,10 ),line_color='red',show=False )
plot_x = sympy.plotting.plot( x.subs( a0,-0.5 ),( t,0,10 ),show=False )
plot_v.append( plot_x[ 0 ] )
plot_v.show()
```